# Strings

Topics:
1. How to get a string
2. String interpolation
3. String concatenation

## How to get a string

Enclose your characters in " " or """ """!

```julia
s1 = "I am a string."
```

Output:
```
"I am a string."
```

```julia
s2 = """I am also a string. """
```

Output:
```
"I am also a string. "
```

There are a couple functional differences between strings enclosed in single and triple quotes. <br>
One difference is that, in the latter case, you can use quotation marks within your string.

```julia
"Here, we get an "error" because it's ambiguous where this string ends "
```

Output:
```
@ ./boot.jl:360 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1116
```

```julia
#use this if we want to use quotation marks within the string
"""Look, Mom, no "errors"!!! """
```

Output:
```
"Look, Mom, no \"errors\"!!! "
```

Note that ' ' define a character, but NOT a string!

```julia
#this is how we create a character variable but NOT a string
typeof('a')
```

Output:
```
Char
```

```julia
#this is not a character but a proper sentence/string variable we would like to declare
'We will get an error here'
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[1]:2
```

## String interpolation

We can use the $ sign to insert existing variables into a string and to evaluate expressions within a string. <br>
Below is an example that contains some highly sensitive personal information.

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
```

Output:
```
10
```

```julia
#we can use the dollar sign to replace with the actual variable, e.g. name, num_fingers and num_toes
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

Output:
```
Hello, my name is Jane.
I have 10 fingers and 10 toes.
```

```julia
# we can also use this to perform calculations within a string! 
println("That is $(num_fingers + num_toes) digits in all!!")
```

Output:
```
That is 20 digits in all!!
```

## String concatenation

Below are three ways we can concatenate strings! <br>
#### 1) Using the string function to connect strings
The first way is to use the `string()` function. <br>
`string()` converts non-string inputs to strings.

```julia
# we can use string to connet two string.
string("How many cats ", "are too many cats?")
```

Output:
```
"How many cats are too many cats?"
```

```julia
# we can even add integers in between the strings
string("I do not know, but ", 10, " are too few.")
```

Output:
```
"I do not know, but 10 are too few."
```

#### 2) Storing strings in variables and using string()

```julia
# we can store the strings stored in variables
s3 = "How many cats ";
s4 = "is too many cats?";
😺 = 10
```

Output:
```
10
```

```julia
string(s3, s4)
```

Output:
```
"How many cats is too many cats?"
```

```julia
string("I don't know, but ", 😺, " is too few.")
```

Output:
```
"I don't know, but 10 is too few."
```

#### 3) Storing strings in variables and using *
We can also use `*` for concatenation!

```julia
s3*s4
```

Output:
```
"How many cats is too many cats?"
```

### Exercises

#### 2.1 
Create a string that says "hi" 1000 times, first with `repeat` and then with the exponentiation operator, which can call `*` under the hood. Assign it the variable `hi` below.

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia
input = "hi"
# find the repeat function here: https://docs.julialang.org/en/v1/base/arrays/#Base.repeat
hi = repeat(input::AbstractString, 1000::Integer)
```

Output:
```
"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"[93m[1m ⋯ 1460 bytes ⋯ [22m[39m"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

<!-- grade_id: cell-20b1027a968fbb48, points: 1, grade: true, solution: false -->
```julia
@assert hi == "hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

```julia
#now, let's create it with the exponentation operator 

hi_2 = input ^ (10* 100)
```

Output:
```
"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"[93m[1m ⋯ 1460 bytes ⋯ [22m[39m"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

```julia
@assert hi == hi_2
```

#### 2.2 
Declare two variables

```julia
a = 3
b = 4
```
and use them to create two strings:
```julia
"3 + 4"
"7" 
```
and store the results in `c` and `d` respectively

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia
a = 3
b = 4
```

Output:
```
4
```

```julia
c = string(a, " + ", b)
d = string("$( a + b)")

#just double-checking before asserting;)
println(c)
print(d)
```

Output:
```
3 + 4
7
```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3 + 4"
@assert d == "7"
```